# Homework 2: DataFrames, Data Visualization, and Functions

## Due Wednesday, January 28th at 11:59PM

Welcome to Homework 2! This week, we will cover DataFrame manipulations, making visualizations, and defining functions. You can find additional help on these topics in  [BPD 6, 9-12](https://notes.dsc10.com/01-getting_started/functions-defining.html) in the `babypandas` notes and [CIT 7-7.3](https://inferentialthinking.com/chapters/07/Visualization.html) in the textbook.

### Instructions

Remember to start early and submit often. You are given six slip days throughout the quarter to extend deadlines. See the syllabus for more details. With the exception of using slip days, late work will not be accepted unless you have made special arrangements with your instructor.

**Important**: For homeworks, the `otter` tests don't usually tell you that your answer is correct. More often, they help catch careless mistakes. It's up to you to ensure that your answer is correct. If you're not sure, ask someone (not for the answer, but for some guidance about your approach). These are great questions for office hours (see the schedule on the [Calendar](https://dsc10.com/calendar)) or Campuswire. Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

**Please do not use for-loops for any questions in this homework.** If you don't know what a for-loop is, don't worry – we haven't covered them yet. But if you do know what they are and are wondering why it's not OK to use them, it is because loops in Python are slow, and looping over arrays and DataFrames should usually be avoided.

<font color=red>**🚨 If you create a data visualization that is too cluttered to read or takes more than a few seconds to generate, this is a sign you are doing something wrong. Do not submit code like this, or the Gradescope autograder may fail to grade your entire assignment. It's better to leave a question blank than to submit code that will cause the autograder to fail on your full assignment.**</font>

```python
# Please don't change this cell, but do make sure to run it
import babypandas as bpd
import numpy as np

import matplotlib.pyplot as plt
plt.style.use('ggplot')
plt.rcParams['figure.figsize'] = (10, 5)

import otter
grader = otter.Notebook()
```

## 1. UCSD ArtPower Events 🎨🎺💃

<center><img src="./images/artpower.jpg" width=600/></center>

As a UCSD student, you may have noticed some emails from <a href="https://artpower.ucsd.edu/">ArtPower</a>, a campus-based performing arts initiative, advertising upcoming events. ArtPower's self-described mission is to "present performing arts that engage, energize, and transform the diverse cultural life of the university and San Diego." ArtPower events are open to the general public, and many are free for UCSD students. They encompass all sorts of performing arts, such as music, dance, and theatre. 

In this question, we'll work with a dataset of upcoming ArtPower events in the 2024-25 season. This dataset was collected by DSC 10 tutor, Austin Flippo, who wrote a <a href="https://github.com/AustinFlippo/UCSD_Events_Scraper">web scraper</a> to collect information from the <a href="https://artpower.ucsd.edu/events/"> ArtPower Events website</a> and store this information in a DataFrame. This DataFrame, `artpower`, is provided below, and it contains the following information for each event. 

|Column|Description|
|------|-----------|
|`'Event Name'`|The name of the event|
|`'Category'`|The type of event|
|`'Description'`|A brief summary of the event|
|`'Link'`|The link to the event webpage|
|`'Price'`|A breakdown of the pricing for different groups of people|
|`'Date'`|Day, month, and year of the event|
|`'Time'`|The day of the week and the time of the event|
|`'Location'`|Location of the event (often on campus, but not always)|
|`'Latitude'`|The latitude for the event's location|
|`'Longitude'`|The longitude for the event's location|

```python
artpower = bpd.read_csv('data/artpower.csv')
artpower
```

Let's start by exploring the `'Event Name'` column using the `describe()` method.

```python
artpower.get('Event Name').describe()
```

**Question 1.1.** From the output above, we see that there are 24 event names, but only 23 of them are unique. This means that there are two events with the same name in `artpower`. Find the name of the event that is duplicated and store the result in `same_name` below.

***Hint:*** Can you organize events so that events with the same name wind up together?

```python
same_name = ...
same_name
```

```python
grader.check("q1_1")
```

**Question 1.2.** Are these actually two different events? Or is there an error in the data whereby we actually have two rows corresponding to the same event at the same time and place? To answer this question, make a new DataFrame, `same_name_events`, which should have the same columns as `artpower`, but only the two rows whose `'Event Name'` is `same_name`, as you found in the previous question. Look to see if there are any differences between the two rows!

```python
same_name_events = ...
same_name_events
```

```python
grader.check("q1_2")
```

Next, let's look at the different event `'Category'` values that appear in the `artpower` DataFrame. You may notice that some events have a `'Category'` of `'Chamber Music'`, while others have a `'Category'` of `'Chamber Music*'`. 

```python
artpower.get('Category').unique()
```

Why the asterisk, you ask? Great question! Here is ArtPower's explanation:

>**ASTERISK\***
>When you see this symbol\*, it indicates artists that are pushing the boundaries, often across genres.

**Question 1.3.** Then create a DataFrame called `chamber_music` that contains all the rows of `artpower` for events where the `'Category'` is `'Chamber Music'`, and all the rows for events where the `'Category'` is `'Chamber Music*'`, but no additional rows. 

```python
chamber_music = ...
chamber_music
```

```python
grader.check("q1_3")
```

**Question 1.4.** Suppose you want to extract the month of the year (e.g. February) for each event. Write a function called `month_of_year` that takes as input one entry from the ``'Date'`` column and outputs the three-letter abbreviation of the month of the year associated with that entry. For example, the function call `month_of_year('18-Feb-25')` should return the string `'Feb'`. You can assume that the inputs to your function will always be in the format shown in the `'Date'` column of `artpower`.

```python
def month_of_year(date):
    ...
```

```python
grader.check("q1_4")
```

**Question 1.5.** Using the `month_of_year` function you just wrote, assign the variable `most_month` to the month of the year that has the *highest* number of events planned. Then, assign `least_month` to the month of the year in the `artpower` DataFrame that has the *lowest* number of events planned. Finally, assign `difference_month` to the difference between the number of events in the most popular (`most_month`) and least popular (`least_month`) months of the year. 

***Hints:*** 
- You will need to define additional variables to answer this question. Feel free to add more cells above the provided cell.
- Start by adding a column to `artpower` that contains the month of the year for each event.

```python
most_month = ...
least_month = ...
difference_month = ...
print("Events occur the most in", most_month)
print("Events occur the least in", least_month)
print(most_month, "has", difference_month, "more events planned than", least_month)
```

```python
grader.check("q1_5")
```

**Question 1.6.** Construct an appropriate visualization to compare the distribution of the number of events across different months of the year, sorted so that the month with the most events appears first in your visualization.

***Notes:*** 
- With this problem, and with all others, make sure your solutions sticks to the methods of the course. While there are many ways to make plots in Python, we've only used one, `.plot`.
- This problem has no grader check. It will be manually graded by our staff.

```python
# Create your plot here.
...
```

```python
# Create your plot here.
...
```

```python
import folium
from folium.plugins import MarkerCluster

def create_map():
    my_map = folium.Map(location=[32.8, -117.2], zoom_start=11)
    marker_cluster = MarkerCluster().add_to(my_map)
    for i in np.arange(artpower.shape[0]):
        folium.Marker(
            location=[artpower.get('Latitude').iloc[i], artpower.get('Longitude').iloc[i]],
            tooltiap="Click me!",
            popup=artpower.get('Event Name').iloc[i] + ' on ' + artpower.get('Date').iloc[i],
            icon=folium.Icon(icon="star")
        ).add_to(marker_cluster)
    return my_map

create_map()
```

```python
on_campus = ...
off_campus = ...
print("There are", on_campus, "events on campus and", off_campus, "events off campus.")
```

```python
grader.check("q1_8")
```

## 2. Broadway Bound 🎭

<center><img src='./images/hamilton.jpg' width = 400/></center>

The two-part film series *Wicked*, a major box office success, is based on the Broadway musical of the same name. In this problem, we'll analyze other Broadway shows. We've provided a file, `./data/broadway.csv` that contains information about Broadway plays and musicals from 1985 to 2020. The dataset comes from [Kaggle](https://www.kaggle.com/datasets/jessemostipak/broadway-weekly-grosses?select=grosses.csv) and was originally scraped from [Playbill](https://playbill.com/grosses) contains a row for each week for each show on Broadway. For each week and show combination, we have the following pieces of information that we will use to generate insights about the state of musicals over time. 

|Column|Description|
|------|-----------|
|`'Week Ending'`|The week, referenced by the Sunday at the end of the week.|
|`'Show'`|The name of the play or musical.|
|`'Theatre'`|The location or venue at which the show was held.|
|`'Weekly Gross'`|The total revenue generated from ticket sales for that show during that week.|
|`'Seats Sold'`|The total number of seats sold for that show during that week.|
|`'Performances'`|The number of performances of the show held during that week, including "previews" before the official opening.|

```python
# Run this cell to load the dataset.
broadway = bpd.read_csv('data/broadway.csv') 
broadway
```

**Question 2.1.** Find the name of the show that generated the most revenue during the timespan of our dataset. Assign your answer to `most_revenue_show`. 

```python
most_revenue_show = ...
most_revenue_show
```

```python
grader.check("q2_1")
```

**Question 2.2.** Broadway shows come and go, but some stay longer than others. Which show in `broadway` was on Broadway for the greatest number of weeks (not necessarily consecutive)? Assign you answer to `most_weeks_show`.

```python
most_weeks_show = ...
most_weeks_show
```

```python
grader.check("q2_2")
```

**Question 2.3.** Add a column to `broadway` that contains, for each week and show combination, the average price of a ticket for that show during that week. Call your new column `'AvgTicket'`, and save the resulting DataFrame in the variable `broadway`.

```python
broadway = ...
broadway                    
```

```python
grader.check("q2_3")
```

**Question 2.4.** The musical *Wicked* has been at the Gershwin Theatre continuously since its opening in 2003. Create a line plot that displays how the average weekly ticket price for *Wicked* has changed over time.

```python
# Create your line plot here.
...
```

```python
with_month = ...
with_month
```

```python
grader.check("q2_5")
```

**Question 2.6.** Create a line plot showing the mean value of `'AvgTicket'` for each month of the year, to see the seasonal variation. Then, assign `most_expensive` to the month of the year (again, as a string like `'04'`) where ticket prices tend to be the most expensive. Extract this month using Python code, but verify visually that it matches up with what your line plot shows.

```python
# Create your line plot here.
with_month.get(['Month', 'AvgTicket']).groupby('Month').mean().plot(kind='line', y='AvgTicket'); # SOLUTION 
```

```python
most_expensive = ...
most_expensive 
```

```python
grader.check("q2_6")
```

**Question 2.7.** Which Broadway theatres have hosted the most performances? To answer this question, create a DataFrame named `theatre_proportions` indexed by `Theatre`, with one column `'Proportion'`. The `'Proportion'` column should contain, for each theatre, the proportion of performances in the dataset that were held in that theatre. Remember that each row of our dataset corresponds to multiple performances! Order the rows of `theatre_proportions` in descending order of `'Proportion'`. 

```python
theatre_proportions = ...
theatre_proportions
```

```python
grader.check("q2_7")
```

**Question 2.8.** Broadway shows tend to bring in a lot of money each week! With multiple performances, packed theatres, and high ticket prices, the revenue stream is flowing! Let's explore how much money was brought in each week by the show *Hamilton*,  a popular award-winning musical that debuted on Broadway in 2015. Plot a density histogram that shows the distribution of `'Weekly Gross'` for *Hamilton* over the period between its debut and the end of data collection. Use bins with a width of a quarter of a million dollars, and make sure that all the data from *Hamilton* is included within your bins.

```python
# Create your histogram here.
...
```

```python
prop_between = ...
prop_between
```

```python
grader.check("q2_9")
```

**Question 2.10.** Broadway shows often have interesting and memorable names that reflect their themes, characters, or settings. From classics like *The Phantom of the Opera* to modern hits like *Hamilton*, these names capture the imagination and hint at the unique stories waiting to be told on stage. How often do Broadway shows get one-word names? Assign `one_word` to the proportion of Broadway show names that are only one word. Remember, each show appears multiple times in our DataFrame!

***Note:*** Some shows are made up of numbers only, which are not really "words" in the conventional sense. We'll consider words to be chunks of text separated by spaces, so for example, the show `"1776"` should count as having one word.

```python
one_word = ...
one_word
```

```python
grader.check("q2_10")
```

## 3. Final Stretch 🧘‍♀️

Suppose we have a DataFrame named `data` with two numerical columns, `'x'` and `'y'`. Consider the following scatter plot, which was generated by calling `data.plot(kind='scatter', x='x', y='y')`:

<img src='images/q3_scatter.png' width=400/>

Now consider these two histograms:

<center>
    <table><tr>
        <td><center><b>Histogram A</b><br> <img src='images/q3_hist_one.png' width=400></center> </td>
        <td><center><b>Histogram B</b><br> <img src='images/q3_hist_two.png' width=400></center> </td>
    </tr></table>
</center>

**Question 3.1.** Which of the following lines of code generated Histogram A? Assign either `1`, `2`, `3`, or `4` to `which_code`.

 1. `data.plot(kind='hist', density=False, y='x')`
 2. `data.plot(kind='hist', density=False, y='y')` 
 3. `data.plot(kind='hist', density=True, y='x')`
 4. `data.plot(kind='hist', density=True, y='y')`

```python
which_code = ...
```

```python
grader.check("q3_1")
```

**Question 3.2.** Suppose we run this block of code:

```py
new_data = bpd.DataFrame().assign(
    x = data.get('x') / 5,
    y = data.get('y')
)
```
    
We then run 

`new_data.plot(kind='hist', density=True, y='x')`.

How will this histogram look compared to the histogram created by 

`data.plot(kind='hist', density=True, y='x')`, 

assuming both histograms are drawn on the same axes? Assign `histogram_difference` to either 1, 2, 3, or 4, corresponding to your choice.

1. The `new_data` histogram will be wider and taller than the `data` histogram.
2. The `new_data` histogram will be wider and shorter than the `data` histogram.
3. The `new_data` histogram will be narrower and taller than the `data` histogram.
4. The `new_data` histogram will be narrower and shorter than the `data` histogram.

_*Hint*_: Look at the end of [Lecture 7](https://dsc10.com/resources/lectures/lec07/lec07.html#Plotting-overlaid-histograms) for an example of two histograms drawn on the same axes.

```python
histogram_difference = ...
```

```python
grader.check("q3_2")
```

**Question 3.3.** Below, we show Histogram B again.

<img src="./images/q3_hist_two.png" width=400/>

What **percent** of values in Histogram B are between -4 (inclusive) and -2 (exclusive)? While we cannot answer this question exactly since we do not know where the bins start and end, we can still approximate the answer. Assign the variable `percent_between` to a number 1 through 5, corresponding to the closest answer.

1. 10% 
2. 13% 
3. 27%
4. 35%
5. 48%

```python
percent_between = ...
```

```python
grader.check("q3_3")
```

## Finish Line: Almost there, but make sure to follow the steps below to submit! 🏁

**_Citations:_** Did you use any generative artificial intelligence tools to assist you on this assignment? If so, please state, for each tool you used, the name of the tool (ex. ChatGPT) and the problem(s) in this assignment where you used the tool for help.

<hr style='color:Maroon;background-color:Maroon;border:0 none; height: 3px;'>

Please cite tools here.

<hr style='color:Maroon;background-color:Maroon;border:0 none; height: 3px;'>

To submit your assignment:

1. Select `Kernel -> Restart & Run All` to ensure that you have executed all cells, including the test cells. 
1. Read through the notebook to make sure all cells ran and all tests passed.
1. Run the cell below to run all tests, and make sure that they all pass.
1. Download your notebook using `File -> Download as -> Notebook (.ipynb)`, then upload your notebook to Gradescope.
1. Stick around while the Gradescope autograder grades your work. Make sure you see that all tests have passed on Gradescope.
1. Check that you have a confirmation email from Gradescope and save it as proof of your submission. 

With homeworks, unlike with labs, the grade you see on Gradescope is **not your final score**. We will run correctness tests after the assignment's due date has passed.

```python
grader.check_all()
```

```python

```

Tests for q1_1 (points: None)
```
>>> isinstance(same_name, str)
True
```


Tests for q1_2 (points: None)
```
>>> isinstance(same_name_events, bpd.DataFrame)
True
```
```
>>> same_name_events.shape == (2, 10)
True
```


Tests for q1_3 (points: None)
```
>>> isinstance(chamber_music, bpd.DataFrame)
True
```


Tests for q1_4 (points: None)
```
>>> isinstance(month_of_year('18-Feb-25'), str)
True
```
```
>>> month_of_year('18-Feb-25').isalpha()
True
```
failure_message: Make sure to exclude the year and day from your answer!
```
>>> len(month_of_year('18-Feb-25')) == 3
True
```


Tests for q1_5 (points: None)
```
>>> isinstance(most_month, str) and isinstance(least_month, str)
True
```
```
>>> import numbers
>>> isinstance(difference_month, numbers.Integral)
True
```
```
>>> bool(difference_month >= 0)
True
```
failure_message: Subtract in the other order.


Tests for q1_8 (points: None)
```
>>> import numbers
>>> isinstance(on_campus, numbers.Integral) and isinstance(off_campus, numbers.Integral)
True
```
```
>>> on_campus + off_campus == artpower.shape[0]
True
```


Tests for q2_1 (points: None)
```
>>> most_revenue_show in broadway.get('Show').unique()
True
```
failure_message: most_revenue_show should be the name of a show


Tests for q2_10 (points: None)
```
>>> isinstance(one_word, float)
True
```


Tests for q2_2 (points: None)
```
>>> isinstance(most_weeks_show, str)
True
```
```
>>> most_weeks_show in broadway.get('Show').unique()
True
```
failure_message: most_weeks_show should be the name of a show


Tests for q2_3 (points: None)
```
>>> broadway.shape[1] == 7
True
```
```
>>> 'AvgTicket' in broadway.columns
True
```
failure_message: Check capitalization and spelling of your column name.


Tests for q2_5 (points: None)
```
>>> broadway.shape[1] == 7 and with_month.shape[1] == 8 and ('Month' in with_month.columns)
True
```


Tests for q2_6 (points: None)
```
>>> most_expensive in ['01', '02', '03', '04', '05', '06', '07', '08', '09', '10', '11', '12']
True
```


Tests for q2_7 (points: None)
```
>>> isinstance(theatre_proportions, bpd.DataFrame) and set(theatre_proportions.columns) == {'Proportion'}
True
```
```
>>> bool(np.isclose(theatre_proportions.get('Proportion').sum(), 1))
True
```
```
>>> theatre_proportions.index[1] == 'Minskoff Theatre'
True
```
failure_message: Tests that the DataFrame is sorted.


Tests for q2_9 (points: None)
```
>>> isinstance(prop_between, float)
True
```
```
>>> prop_between >= 0 and prop_between <= 1
True
```


Tests for q3_1 (points: None)
```
>>> which_code in {1, 2, 3, 4}
True
```


Tests for q3_2 (points: None)
```
>>> histogram_difference in {1, 2, 3, 4}
True
```


Tests for q3_3 (points: None)
```
>>> percent_between in {1, 2, 3, 4, 5}
True
```
