# Homework 6: Hypothesis Testing and Permutation Testing

## Due Thursday, March 5th at 11:59PM

Welcome to Homework 6, the last homework of the quarter! This homework covers hypothesis testing ([CIT 11](https://inferentialthinking.com/chapters/11/Testing_Hypotheses.html)) and permutation testing ([CIT 12](https://inferentialthinking.com/chapters/12/Comparing_Two_Samples.html)).

### Instructions

You are given six slip days throughout the quarter to extend deadlines. See the syllabus for more details. With the exception of using slip days, late work will not be accepted unless you have made special arrangements with your instructor.

**Important**: For homeworks, the `otter` tests don't usually tell you that your answer is correct. More often, they help catch careless mistakes. It's up to you to ensure that your answer is correct. If you're not sure, ask someone (not for the answer, but for some guidance about your approach). These are great questions for office hours (see the schedule on the [Calendar](https://dsc10.com/calendar)) or Campuswire. Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged.

```python
# Please don't change this cell, but do make sure to run it
import babypandas as bpd
import numpy as np
np.set_printoptions(legacy='1.21')

import matplotlib.pyplot as plt
plt.style.use('seaborn-v0_8-colorblind')

import otter
grader = otter.Notebook()
```

## 1. Secret Smiski 🧍

<center><img src='images/smiskis.png' width='600'></center>

Smiskis are tiny glow-in-the-dark collectible figurines. Each Smiski is packaged in an identical box so that you don't know which Smiski you have when you pick up a box. Due to this packaging, it isn't easy to get a specific Smiski, but it's even rarer to get a special type of Smiski called a *secret* Smiski. Although no one knows the exact probability of obtaining a secret Smiski, you [read online](https://www.reddit.com/r/smiskis/comments/15myvll/how_rare_are_secret_smiskis_really/) that chance of obtaining a secret Smiski is $\frac{1}{144}$. 

Many of the DSC 10 tutors are avid Smiski collectors and have collectively bought 903 Smiskis over the past few years. Out of all the Smiskis they collected, they got 10 secret Smiskis. Due to this outcome, the tutors suspect the probability of getting a secret Smiski should be higher than $\frac{1}{144}$.

To test this, they decide to run a hypothesis test with the following hypotheses:

**Null Hypothesis**: The probability of obtaining a secret Smiski is $\frac{1}{144}$. 

**Alternative Hypothesis**: The probability of obtaining a secret Smiski is greater than $\frac{1}{144}$.

**Question 1.1.** Complete the implementation of the function `one_simulation`, which has no arguments. It should randomly generate 903 Smiskis under the assumptions of the null hypothesis and return the **proportion** of Smiskis that are secret Smiskis. 

***Hint:*** Use `np.random.multinomial`. You don't need a `for`-loop.

```python
def one_simulation():
    ...
```

```python
grader.check("q1_1")
```

The test statistic for our hypothesis test will be the difference between the proportion of secret Smiskis obtained in a given sample of 903 Smiskis and the expected proportion of secret Smiskis. 

$$\text{test statistic} = \text{proportion of secret Smiskis in sample} - \frac{1}{144}$$

**Question 1.2.** Calculate the observed value of the test statistic and store the result in `smiski_observed`. Recall that the DSC 10 tutors have 903 Smiskis, 10 of which are secret Smiskis.

```python
smiski_observed = ...
smiski_observed
```

```python
grader.check("q1_2")
```

**Question 1.3.** Let's conduct 10,000 simulations. Create an array named `proportion_diffs` containing 10,000 simulated values of the test statistic described above. Utilize the function created in the previous question to perform this task.

```python
proportion_diffs = ...

# Visualize with a histogram. Don't change anything below.
bpd.DataFrame().assign(proportion_differences=proportion_diffs).plot(kind='hist', bins=15, density=True, ec='w', figsize=(10, 5));
plt.axvline(x=smiski_observed, color='black', linewidth=4, label='observed statistic')
plt.legend();
```

```python
grader.check("q1_3")
```

**Question 1.4.** Calculate the p-value for this hypothesis test, and assign the result to `smiski_p`.

***Hint:*** Do large values of our test statistic favor the alternative hypothesis, or do small values of our test statistic favor the alternative hypothesis?

```python
smiski_p = ...
smiski_p
```

```python
grader.check("q1_4")
```

**Question 1.5.** Using the standard p-value cutoff of 0.05, what can we conclude from our hypothesis test? Assign either 1, 2, 3, or 4 to the variable `smiski_conclusion`, corresponding to the best conclusion.

   1. We reject the null hypothesis. There is not enough evidence to draw a conclusion about whether the data is consistent with the model.
   1. We reject the null hypothesis. The observed data is inconsistent with the model.
   1. We accept the null hypothesis. The observed data is consistent with the model.
   1. We fail to reject the null hypothesis. There is not enough evidence to say that the observed data is inconsistent with the model.

```python
smiski_conclusion = ...
```

```python
grader.check("q1_5")
```

**Question 1.6.** In this question, we chose as our test statistic the proportion of secret Smiskis obtained minus $\frac{1}{144}$. But this is not the only statistic we could have chosen; there are many that could have worked here. 

From the options below, choose the test statistics that would **not** have worked for this hypothesis test, and assign the variable `bad_choices` to a list of your choices.

1. The number of secret Smiskis obtained out of 903 Smiskis.
1. The absolute difference between 10 and the number of secret Smiskis obtained.
1. The absolute difference between $\frac{1}{144}$ and the proportion of secret Smiskis obtained.
1. $\frac{1}{144}$ minus the proportion of secret Smiskis obtained.
1. The proportion of secret Smiskis obtained.

***Hint:*** Our goal is to find a test statistic that will help us determine whether we got secret Smiskis **more** often than expected.

```python
bad_choices = ...
bad_choices
```

```python
grader.check("q1_6")
```

## 2. Gotta Catch 'Em All ✅

<center><img src='./images/pokemon.jpg' width='400'></center>

[Pokémon Trading Card Game Pocket (TCGP)](https://tcgpocket.pokemon.com/en-us/) is a mobile adaptation of the classic Pokémon trading card game. Players build decks using cards that feature different Pokémon (short for Pocket Monsters). 

There are different tiers for Pokémon cards, representing their rarity. Diamonds are used for the common cards, with one diamond being the most common. Cards with a star are more rare, and finally, crowns represent the rarest cards. 

Eric is a DSC 10 tutor who has been playing TCGP for a while. He wonders how cards are distributed across the different rarity tiers. Based on his extensive gameplay, he proposes the following probability distribution for how frequently each type of card appears. Note that the sum of the estimated probabilities is 1.

| Rarity Tier | Eric's Estimated Probability|
| --- | --- |
| One Diamond | $0.25$ |
| Two Diamonds | $0.2$ |
| Three Diamonds | $0.15$ |
| Four Diamonds | $0.1$ |
| One Star| $0.09$ |
| Two Stars | $0.08$ |
| Three Stars| $0.07$ |
| Crown |$0.06$|

We'll store this **proposed** distribution in an array, in the order shown above.

```python
# Just run this cell, do not change it!
proposed_dist = np.array([0.25, 0.2, 0.15, 0.1, 0.09, 0.08, 0.07, 0.06])
proposed_dist
```

To assess the validity of Eric's model, you collect data directly from Pokémon TCGP. You learn that the last 1,000 cards were as follows:

| Rarity Tier | Number of Cards|
| --- | --- |
| One Diamond | $264$ |
| Two Diamonds | $224$ |
| Three Diamonds | $154$ |
| Four Diamonds | $105$ |
| One Star| $88$ |
| Two Stars | $68$ |
| Three Stars| $62$ |
| Crown |$35$|

You then calculate the **observed** distribution using the data you collected and store it in an array as well (in the same order as before):

```python
# Just run this cell, do not change it!
observed_dist = np.array([264, 224, 154, 105, 88, 68, 62, 35]) / 1000
observed_dist
```

While `observed_dist` is not identical to `proposed_dist`, it's still possible that Eric's model is plausible, and that the differences are simply due to random chance. Let's run a hypothesis test to investigate further, using the following hypotheses: 

- **Null Hypothesis**: Pokémon TCGP Cards are randomly drawn from the distribution `proposed_dist`.

- **Alternative Hypothesis**: Pokémon TCGP Cards are _not_ drawn randomly from the distribution `proposed_dist`.

Note that this hypothesis test involves eight proportions, one for each rarity tier.

**Question 2.1.**  Which of the following is **not** a reasonable choice of test statistic for this hypothesis test? Assign 1, 2, or 3 to the variable `unreasonable_test_statistic`. 
1. The sum of the absolute differences between the proposed distribution (Eric's expected proportion of rarities) and the observed distribution (actual proportion of rarities).
1. The absolute difference between the sum of the proposed distribution (Eric's expected proportion of rarities) and the sum of the observed distribution (actual proportion of rarities).
1. Among all eight card rarities, the largest absolute difference between Eric's expected proportion and the actual proportion of cards of that rarity.

```python
unreasonable_test_statistic = ...
```

```python
grader.check("q2_1")
```

**Question 2.2.** We'll use the TVD, i.e. **total variation distance**, as our test statistic. Below, complete the implementation of the function `total_variation_distance`, which takes as input two distributions (stored as arrays) and returns the total variation distance between those distributions.

Then, use the function `total_variation_distance` to determine the TVD between the distribution proposed by Eric and the observed distribution of rarities. Assign this TVD to `observed_tvd`.

```python
def total_variation_distance(first_distrib, second_distrib):
    '''Computes the total variation distance between two distributions.'''
    ...

observed_tvd = ...
observed_tvd
```

```python
grader.check("q2_2")
```

**Question 2.3.** Now, we'll calculate 5,000 simulated TVDs to see what a typical TVD between the proposed distribution and a simulated distribution would look like if Eric's model were accurate. Since our real-life data includes 1000 cards, in each trial of the simulation, we'll:
- draw 1000 cards at random from Eric's proposed distribution, then 
- calculate the TVD between **Eric's proposed type distribution** and the **type distribution from the simulated sample**. 

Store these 5,000 simulated TVDs in an array called `simulated_tvds`.

```python
simulated_tvds = ...

# Visualize the distribution of TVDs with a histogram
bpd.DataFrame().assign(simulated_tvds=simulated_tvds).plot(kind='hist', density=True, ec='w', figsize=(10, 5));
plt.axvline(x=observed_tvd, color='black', linewidth=4, label='observed TVD')
plt.legend();
```

```python
grader.check("q2_3")
```

**Question 2.4.** Now, determine the p-value for our test by finding the proportion of times in our simulation that we saw a TVD greater than or equal to our observed TVD. Assign your result to `pokemon_p`.

```python
pokemon_p = ...
pokemon_p
```

```python
grader.check("q2_4")
```

**Question 2.5.** Using a p-value cutoff of 0.01, what can we conclude from our hypothesis test? Assign either 1, 2, 3, or 4 to the variable `pokemon_conclusion`, corresponding to the best conclusion.
   
   1. We accept the null hypothesis. The observed data is consistent with the model.
   1. We reject the null hypothesis. There is not enough evidence to say that the observed data is consistent with the model.
   1. We reject the null hypothesis. The observed data is inconsistent with the model.
   1. We fail to reject the null hypothesis. There is not enough evidence to say that the observed data is inconsistent with the model.

```python
pokemon_conclusion = ...
pokemon_conclusion
```

```python
grader.check("q2_5")
```

## 3. Chocolate 🍫😋
<img src='images/chocolate_bars.png' width='1000'>

Chocolate is a well-loved treat that many enjoy, but some people take their chocolate very seriously. [The Manhattan Chocolate Society](https://flavorsofcacao.com/mcs_index.html) is an invitation-only society founded to taste and review dark chocolate bars from around the world. The [Flavors of Cacao database](https://flavorsofcacao.com/index.html) was born from tastings done by this exclusive society, and it contains reviews of nearly three thousand different dark chocolate bars. Which dark chocolate bars do these connoisseurs consider to be the best? Let's find out!

Run the next cell to load in the data.

```python
choco = bpd.read_csv('data/chocolate.csv')
choco
```

We will primarily be working with the `'Characteristics'` and `'Rating'` columns. The `'Rating'` column contains a score from 1 to 5. According to Flavors of Cacao, each rating can be interpreted as follows:

| Rating | Meaning |
| ------ | ------- |
| 4.0 - 5.0  | Outstanding |
| 3.5 - 3.9  | Highly Recommended |
| 3.0 - 3.49 | Recommended |
| 2.0 - 2.9  | Disappointing |
| 1.0 - 1.9  | Unpleasant |

Ratings are determined by a combination of factors including flavor, texture, and "aftermelt", or the lingering experience after the chocolate has melted in your mouth.

The `'Characteristics'` column contains the *most memorable characteristics* of each dark chocolate bar. Each bar may have several memorable characteristics, separated by a comma. For example, the chocolate bar at the last index of the DataFrame was memorable for its woody flavor and butterscotch notes.

Compared to other types of chocolate, dark chocolate tends to be less sweet. However, quite a few of the dark chocolate bars in the DataFrame above were memorable for being sweet. How do sweet dark chocolate bars get rated relative to non-sweet dark chocolate bars? In this problem, we will explore whether the ratings for sweet dark chocolate bars come from the same distribution as non-sweet dark chocolate bars. 

**Question 3.1.** Complete the implementation of the function `label_sweet`, which takes in a string of characteristics associated with a single row of `choco` and returns one of two strings, either `'Sweet'` or `'Not Sweet'`. 

We will consider a chocolate bar to be `'Sweet'` if the word `'sweet'` appears in the string of characteristics, but **not** if it appears only as part of a longer word, such as `'bittersweet'`. Example behavior is given below.

```
>>>label_sweet('sweet, woody')
'Sweet'

>>>label_sweet('creamy, too sweet, sour')
'Sweet'

>>>label_sweet('nutty, bittersweet, chalky')
'Not Sweet'
```

```python
def label_sweet(characteristics): 
    ...
```

```python
grader.check("q3_1")
```

**Question 3.2.**  Assign `chocolate` to a DataFrame with a row for each chocolate bar represented in `choco`, but with only two columns, `'Rating'` and `'Sweetness'`. `'Rating'` should be exactly as it appears in `choco`, and `'Sweetness'` should contain one of two distinct values: `'Sweet'` or `'Not Sweet'`, interpreted as we've defined in the previous question. 

```python
chocolate = ...
chocolate
```

```python
grader.check("q3_2")
```

**Question 3.3.** Using the DataFrame `chocolate`, calculate the difference between the **mean** `'Rating'` of sweet dark chocolate bars and non-sweet dark chocolate bars. Assign your answer to `observed_difference`.

$$\text{observed difference} = \text{mean rating of sweet dark chocolate bars} - \text{mean rating of non-sweet dark chocolate bars}$$

```python
observed_difference = ...
observed_difference
```

```python
grader.check("q3_3")
```

**Question 3.4.** What does the number you obtained for `observed_difference` mean? Assign `interpretation` to 1, 2, 3, 4, 5 or 6 corresponding to the best explanation below.

1. In our sample, the mean rating for sweet bars is higher than the mean rating for non-sweet bars by about 16 percent.
1. In our sample, the mean rating for sweet bars is higher than the mean rating for non-sweet bars by about 0.16 percent.
1. In our sample, the mean rating for sweet bars is higher than the mean rating for non-sweet bars by about 0.16 rating points.
1. In our sample, the mean rating for sweet bars is lower than the mean rating for non-sweet bars by about 16 percent.
1. In our sample, the mean rating for sweet bars is lower than the mean rating for non-sweet bars by about 0.16 percent.
1. In our sample, the mean rating for sweet bars is lower than the mean rating for non-sweet bars by about 0.16 rating points.


```python
interpretation = ...
```

```python
grader.check("q3_4")
```

**Question 3.5.** Now we want to conduct a **permutation test** to see if what we observed in our sample is reflective of the larger population of dark chocolate bars.

- **Null Hypothesis**: The ratings of sweet dark chocolate bars and non-sweet dark chocolate bars come from the same distribution.  
- **Alternative Hypothesis**: The ratings of sweet dark chocolate bars are lower on average than the ratings of non-sweet dark chocolate bars.

Run a permutation test to see if the `observed_difference` you calculated in Question 3.3 is actually a statistically significant difference. Simulate 1000 values of the test statistic by shuffling the `'Sweetness'` column of `chocolate` and calculating the difference in mean rating between the two groups determined by the shuffling (again, in the order sweet minus non-sweet). Store your 1000 differences in the `differences` array. 

***Hint:*** It's a good idea to simulate one value of the test statistic before putting everything in a for-loop.

```python
differences = ...

# Just display the first ten differences.
differences[:10]
```

```python
grader.check("q3_5")
```

**Question 3.6.** Compute a p-value for this hypothesis test and assign your answer to `chocolate_p`. To decide whether to use `<=` or `>=` in the calculation of the p-value, think about whether larger values or smaller values of our test statistic favor the alternative hypothesis.

```python
chocolate_p = ...
chocolate_p
```

```python
grader.check("q3_6")
```

**Question 3.7.** Assign the variable `chocolate_conclusion` to a **list** of all the true statements below.

1. We accept the null hypothesis at the 0.01 significance level.
1. We reject the null hypothesis at the 0.01 significance level.
1. We fail to reject the null hypothesis at the 0.01 significance level.
1. We accept the null hypothesis at the 0.05 significance level.
1. We reject the null hypothesis at the 0.05 significance level.
1. We fail to reject the null hypothesis at the 0.05 significance level.

Then, interpret your results by setting `sweeter_is_worse` to `True` or `False`, based on the outcome of your permutation test. `True` means that sweet dark chocolate bars actually do have lower ratings than non-sweet bars, and `False` means they do not.

```python
chocolate_conclusion = ...
sweeter_is_worse = ...
```

```python
grader.check("q3_7")
```

**Question 3.8.** Suppose in this question you had shuffled the `'Rating'` column instead and kept the `'Sweetness'` column in the same order. Assign `shuffled_rating` to either 1, 2, 3, or 4, corresponding to the true statement below.


1. The new p-value from shuffling `'Rating'` would be $1 - p$, where $p$ is the old p-value from shuffling `'Sweetness'` (i.e. your answer to Question 3.6).
1. We would need to change our null hypothesis in order to shuffle the `'Rating'` column. 
1. We would need to change out test statistic in order to shuffle the `'Rating'` column. 
1. There would be no difference in the conclusion of the test if we had shuffled the `'Rating'` column instead.
1. The `'Rating'` column cannot be shuffled because it contains numbers.

```python
shuffled_rating = ...
```

```python
grader.check("q3_8")
```

**Question 3.9.** Which of the following choices best describes the purpose of shuffling one of the columns in our dataset in a permutation test? Assign `why_shuffle` to either 1, 2, 3, or 4, corresponding to the true statement below.

1. Shuffling mitigates noise in our data by generating new permutations of the data.
1. Shuffling is a special case of bootstrapping and allows us to produce interval estimates.
1. Shuffling allows us to generate new data under the null hypothesis, which we can use in testing our hypothesis.
1. Shuffling allows us to generate new data under the alternative hypothesis, which helps us identify when the data come from different distributions.

```python
why_shuffle = ...
```

```python
grader.check("q3_9")
```

Feel free to explore the dark chocolate data some more to see if other characteristics are linked with higher or lower ratings! 

## 4. Battery Life 🔋🎧

You and your friend are studying together for a final. Since both of you focus better when listening to music, you  challenge yourselves to study uninterrupted until each of your earbuds run out of battery. You own a pair of _Apple AirPods_ (i.e., AirPods) and your friend owns a pair of _Samsung Galaxy Buds Pro_ (i.e., Galaxy Buds). You both put on your earbuds and get to work.

At the end of the study session, you find that your AirPods, which were fully charged at the start, lasted 4 hours and 18 minutes, while your friend's Galaxy Buds, also fully charged at the start, lasted 4 hours and 41 minutes. Intrigued by the noticeable difference in battery life between the two earbuds, you decide to investigate further. You make a post on Reddit asking people who have AirPods or Galaxy Buds to record how long their earbuds last on a single charge. You're overwhelmed by the amazing response and receive 80 different comments in total from other people, 40 from people with AirPods and 40 from people with Galaxy Buds.

Let's look at all the data that you crowdsourced. Each entry in the `'BatteryLife'` column represents the amount of time that a pair of earbuds lasted on a full charge, in minutes.

```python
battery_life_data = bpd.read_csv('data/earbuds.csv')
battery_life_data
```

**Question 4.1.** Now let's address the question: how does the average battery life of Galaxy Buds compare to that of AirPods? Create a DataFrame called `galaxy` that contains only the battery life data for Galaxy Buds, and set `galaxy_mean` to the mean battery life of Galaxy Buds. Similarly, create a DataFrame `airpods` for AirPods and compute `airpods_mean`. Finally, set `observed_diff_mean`, to the difference in mean battery life of Galaxy Buds and AirPods in our sample, computed as follows.

$$\text{difference} = \text{mean battery life of Galaxy Buds} - \text{mean battery life of AirPods}$$


```python
galaxy = ...
airpods = ...
galaxy_mean = ...
airpods_mean = ...
observed_diff_mean = ...
observed_diff_mean
```

```python
grader.check("q4_1")
```

If you answered Question 4.1 correctly, you should have noticed a difference in the average battery life between Galaxy Buds and AirPods. But remember, we are only analyzing samples of size 40 for each brand. Would we observe such a difference in battery life if we had access to the entire population – that is, all units ever produced by both brands – or is it possible that this difference is merely a result of the specific samples we happened to collect? Let's do a **hypothesis test** to find out. We'll state our hypotheses as follows:

- **Null Hypothesis**: The average battery life of Galaxy Buds is equal to that of AirPods. In other words, the difference in average battery life between the two brands equals 0 minutes.

- **Alternative Hypothesis**: The average battery life of Galaxy Buds is not equal to that of AirPods. Hence, the difference in average battery life between the two brands is not 0 minutes.


Since we are able to frame our hypothesis test as a question of whether a certain population parameter – the difference in average battery life between Galaxy Buds and AirPods – is equal to a specific value, we can **test our hypotheses by constructing a confidence interval** for this parameter. For a refresher on this method, refer to [CIT 13.4](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html) or the human body temperature example from [Lecture 21](https://dsc10.com/resources/lectures/lec21/lec21.html).

***Note:*** We are **not** conducting a permutation test here, although that would also be a valid approach to test these hypotheses.

**Question 4.2.** Compute 1000 **bootstrapped estimates** for the difference in average battery life between Galaxy Buds and AirPods. As in Question 4.1, calculate the difference as Galaxy Buds minus AirPods. Store your 1000 estimates in the `difference_means` array.

You should generate your Galaxy Buds resamples by sampling from `galaxy`, and your AirPods resamples by sampling from `airpods`. Do not use the combined dataset `battery_life_data` for this task, otherwise you might not wind up with 40 of each!

```python
difference_means = ...

# Just display the first ten differences.
difference_means[:10]
```

```python
grader.check("q4_2")
```

Let's visualize your estimates:

```python
bpd.DataFrame().assign(BootstrappedDifferenceMeans = difference_means).plot(kind = 'hist', density=True, ec='w', bins=20, figsize=(10, 5));
```

**Question 4.3.** Compute a 95% confidence interval for the difference in mean battery life of AirPods and Galaxy Buds (as before, in the order Galaxy Buds minus AirPods). Assign the left and right endpoints of this confidence interval to `left_endpoint` and `right_endpoint` respectively. 

```python
left_endpoint = ...
right_endpoint = ...

print('Bootstrapped 95% confidence interval for the mean difference in battery life of AirPods and Galaxy Buds:\n [{:f}, {:f}]'.format(left_endpoint, right_endpoint))
```

```python
grader.check("q4_3")
```

**Question 4.4.** Based on the confidence interval you've created, would you reject the null hypothesis at the 0.05 significance level? Set `reject_null` to True if you would reject the null hypothesis, and False if you would not.

```python
reject_null = ...
```

```python
grader.check("q4_4")
```

**Question 4.5.** What if all the people who responded to your original Reddit post had provided their battery lives in hours instead of minutes? Would your hypothesis test still come to the same conclusion either way? Set `same_conclusion` to True or False.

```python
same_conclusion = ...
```

```python
grader.check("q4_5")
```

## Finish Line: Almost there, but make sure to follow the steps below to submit! 🏁

**_Citations:_** Did you use any generative artificial intelligence tools to assist you on this assignment? If so, please state, for each tool you used, the name of the tool (ex. ChatGPT) and the problem(s) in this assignment where you used the tool for help.

<hr style="color:Maroon;background-color:Maroon;border:0 none; height: 3px;">

Please cite tools here.

<hr style="color:Maroon;background-color:Maroon;border:0 none; height: 3px;">

Congratulations! You are done with Homework 6 – the final homework of the quarter! 🎉

To submit your assignment:

1. Select `Kernel -> Restart & Run All` to ensure that you have executed all cells, including the test cells.
1. Read through the notebook to make sure everything is fine and all tests passed.
1. Run the cell below to run all tests, and make sure that they all pass.
1. Download your notebook using `File -> Download as -> Notebook (.ipynb)`, then upload your notebook to Gradescope.
1. Stick around while the Gradescope autograder grades your work. Make sure you see that all tests have passed on Gradescope.
1. Check that you have a confirmation email from Gradescope and save it as proof of your submission.

```python
grader.check_all()
```

```python

```

Tests for q1_1 (points: None)
```
>>> callable(one_simulation) and one_simulation() < 0.02
True
```


Tests for q1_2 (points: None)
```
>>> smiski_observed < 0.01
True
```
failure_message: The answer is a very small number.


Tests for q1_3 (points: None)
```
>>> isinstance(proportion_diffs, np.ndarray) and len(proportion_diffs) == 10000
True
```


Tests for q1_4 (points: None)
```
>>> 0 <= smiski_p <= 1
True
```


Tests for q1_5 (points: None)
```
>>> smiski_conclusion in [1, 2, 3, 4]
True
```


Tests for q1_6 (points: None)
```
>>> type(bad_choices) == list and all((choice in [1, 2, 3, 4] for choice in bad_choices))
True
```


Tests for q2_1 (points: None)
```
>>> unreasonable_test_statistic in [1, 2, 3]
True
```


Tests for q2_2 (points: None)
```
>>> callable(total_variation_distance) and 0 <= observed_tvd <= 1
True
```
```
>>> np.isclose(total_variation_distance(np.array([0.5, 0.5]), np.array([1, 0])), 0.5)
True
```
failure_message: Try again!
```
>>> int(observed_tvd * 10 ** 7 % 8 ** 3) == 496
True
```


Tests for q2_3 (points: None)
```
>>> isinstance(simulated_tvds, np.ndarray) and len(simulated_tvds) == 5000
True
```


Tests for q2_4 (points: None)
```
>>> 0 <= pokemon_p <= 1
True
```


Tests for q2_5 (points: None)
```
>>> pokemon_conclusion in [1, 2, 3, 4]
True
```


Tests for q3_1 (points: None)
```
>>> label_sweet('fruit, sweet, nutty') == 'Sweet'
True
```
```
>>> label_sweet('chalky, very sweet, unpleasant') == 'Sweet'
True
```
```
>>> label_sweet('creamy, semisweet, caramel') == 'Not Sweet'
True
```


Tests for q3_2 (points: None)
```
>>> chocolate.shape == (2693, 2) and set(chocolate.columns) == set(['Sweetness', 'Rating'])
True
```
```
>>> np.all(chocolate.get('Sweetness').unique() == np.array(['Not Sweet', 'Sweet']))
True
```
failure_message: Make sure the 'Sweetness' values are either `Sweet` or `Not Sweet`
```
>>> (chocolate.get('Sweetness') == 'Sweet').sum() == 339
True
```
failure_message: You may have an error in your label_sweet function.


Tests for q3_3 (points: None)
```
>>> 'float' in str(observed_difference.dtype) and bool(observed_difference < 0)
True
```
```
>>> int(observed_difference * 10 ** 6 % 10 ** 4) == 8645
True
```
failure_message: Try again!


Tests for q3_4 (points: None)
```
>>> interpretation in [1, 2, 3, 4, 5, 6]
True
```


Tests for q3_5 (points: None)
```
>>> isinstance(differences, np.ndarray) and len(differences) == 1000
True
```


Tests for q3_6 (points: None)
```
>>> isinstance(chocolate_p, float)
True
```


Tests for q3_7 (points: None)
```
>>> type(chocolate_conclusion) == list and type(sweeter_is_worse) == bool
True
```
```
>>> type(chocolate_conclusion) == list and all((choice in [1, 2, 3, 4, 5, 6] for choice in chocolate_conclusion))
True
```


Tests for q3_8 (points: None)
```
>>> shuffled_rating in [1, 2, 3, 4]
True
```


Tests for q3_9 (points: None)
```
>>> why_shuffle in [1, 2, 3, 4]
True
```


Tests for q4_1 (points: None)
```
>>> isinstance(observed_diff_mean, float) and observed_diff_mean > 0
True
```
```
>>> int(observed_diff_mean * 10 ** 2 % 10 ** 2) == 79
True
```
failure_message: Try again!


Tests for q4_2 (points: None)
```
>>> isinstance(difference_means, np.ndarray) and len(difference_means) == 1000
True
```


Tests for q4_3 (points: None)
```
>>> import numbers
>>> isinstance(left_endpoint, numbers.Real) and isinstance(right_endpoint, numbers.Real)
True
```
```
>>> left_endpoint <= right_endpoint
True
```


Tests for q4_4 (points: None)
```
>>> reject_null in [True, False]
True
```


Tests for q4_5 (points: None)
```
>>> same_conclusion in [True, False]
True
```
